# Assignment 1
For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web.



Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    #raise NotImplementedError()
```

## Part A

Find a list of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""

    # YOUR CODE HERE
#     raise NotImplementedError()
    pattern = "[A-Z][a-z]*"
    return re.findall(pattern, simple_string)
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

```

```python
names()
```

Output:
```
['Amy', 'Mary', 'Ruth', 'Peter']
```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in 
a class. Create a regex to generate a list of just those students who received a B in the course.

```python
grades = """
Ronald Mayr: A
Bell Kassulke: B
Jacqueline Rupp: A 
Alexander Zeller: C
Valentina Denk: C 
Simon Loidl: B 
Elias Jovanovic: B 
Stefanie Weninger: A 
Fabian Peer: C 
Hakim Botros: B
Emilie Lorentsen: B
Herman Karlsen: C
Nathalie Delacruz: C
Casey Hartman: C
Lily Walker : A
Gerard Wang: C
Tony Mcdowell: C
Jake Wood: B
Fatemeh Akhtar: B
Kim Weston: B
Nicholas Beatty: A
Kirsten Williams: C
Vaishali Surana: C
Coby Mccormack: C
Yasmin Dar: B
Romy Donnelly: A
Viswamitra Upandhye: B
Kendrick Hilpert: A
Killian Kaufman: B
Elwood Page: B
Mukti Patel: A
Emily Lesch: C
Elodie Booker: B
Jedd Kim: A
Annabel Davies: A
Adnan Chen: B
Jonathan Berg: C
Hank Spinka: B
Agnes Schneider: C
Kimberly Green: A
Lola-Rose Coates: C
Rose Christiansen: C
Shirley Hintz: C
Hannah Bayer: B
"""
```

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re
def grades():
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()

    # YOUR CODE HERE
#     raise NotImplementedError()
    gradeB_pattern = "[\w ]*:\ B"
    gradeB_list = re.findall(gradeB_pattern, grades)
    names_list = []
    for elem in gradeB_list:
        name = elem.split(":")[0]
        names_list.append(name)
    return names_list
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16

```

```python
grades()
```

Output:
```
['Bell Kassulke: B',
 'Simon Loidl: B',
 'Elias Jovanovic: B',
 'Hakim Botros: B',
 'Emilie Lorentsen: B',
 'Jake Wood: B',
 'Fatemeh Akhtar: B',
 'Kim Weston: B',
 'Yasmin Dar: B',
 'Viswamitra Upandhye: B',
 'Killian Kaufman: B',
 'Elwood Page: B',
 'Elodie Booker: B',
 'Adnan Chen: B',
 'Hank Spinka: B',
 'Hannah Bayer: B']
```

```python
def grades():
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()

    # YOUR CODE HERE
#     raise NotImplementedError()
    pattern = "[\w]*\ [\w]*(?=:\ B)"
    return re.findall(pattern, grades)
```

```python
assert len(grades()) == 16
```

```python
grades()
```

Output:
```
['Bell Kassulke',
 'Simon Loidl',
 'Elias Jovanovic',
 'Hakim Botros',
 'Emilie Lorentsen',
 'Jake Wood',
 'Fatemeh Akhtar',
 'Kim Weston',
 'Yasmin Dar',
 'Viswamitra Upandhye',
 'Killian Kaufman',
 'Elwood Page',
 'Elodie Booker',
 'Adnan Chen',
 'Hank Spinka',
 'Hannah Bayer']
```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:
* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re
def logs():
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
    
    # YOUR CODE HERE
#     raise NotImplementedError()
    pattern = """
    (?P<host>[\d]*.[\d]*.[\d]*.[\d]*)    
    (\ -\ )  
    (?P<user_name>[\w-]*) 
    (\ \[) 
    (?P<time>\w*/\w*/.*)
    (\]\ \") 
    (?P<request>.*)
    (")
    """
    # YOUR CODE HERE
    result = []
    for item in re.finditer(pattern, logdata, re.VERBOSE):
        result.append(item.groupdict())
    return result
```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

```

```python
logs()
```

Output:
```
[{'host': '146.204.224.152',
  'request': 'POST /incentivize HTTP/1.1',
  'time': '21/Jun/2019:15:45:24 -0700',
  'user_name': 'feest6811'},
 {'host': '197.109.77.178',
  'request': 'DELETE /virtual/solutions/target/web+services HTTP/2.0',
  'time': '21/Jun/2019:15:45:25 -0700',
  'user_name': 'kertzmann3129'},
 {'host': '156.127.178.177',
  'request': 'DELETE /interactive/transparent/niches/revolutionize HTTP/1.1',
  'time': '21/Jun/2019:15:45:27 -0700',
  'user_name': 'okuneva5222'},
 {'host': '100.32.205.59',
  'request': 'PATCH /architectures HTTP/1.0',
  'time': '21/Jun/2019:15:45:28 -0700',
  'user_name': 'ortiz8891'},
 {'host': '168.95.156.240',
  'request': 'GET /engage HTTP/2.0',
  'time': '21/Jun/2019:15:45:31 -0700',
  'user_name': 'stark2413'},
 {'host': '71.172.239.195',
  'request': 'PUT /cutting-edge HTTP/2.0',
  'time': '21/Jun/2019:15:45:32 -0700',
  'user_name': 'dooley1853'},
 {'host': '180.95.121.94',
  'request': 'PATCH /extensible/reinvent HTTP/1.1',
  'time': '21
```